# Функции

Содержание лекции:
1. Как объявлять функции
2. Утиная типизация в Julia
3. Мутирующие и немутирующие функции
4. Некоторые функции высокого порядка

## Как объявлять функции
Julia предоставляет несколько различных способов для определения функций. Первый способ - с использованием ключевых слов `function` и `end`.

```julia
function sayhi(name)
    println("Привет $name, рад тебя видеть!")
end
```

Output:
```
sayhi (generic function with 1 method)
```

```julia
function f(x)
    return x^2
end
```

Output:
```
f (generic function with 1 method)
```

В дальнейшем эти функции могут быть вызваны следующим образом:

```julia
sayhi("C-3PO")
```

```julia
f(42)
```

Альтернативный способ заключается в определении функии в одну строку

```julia
sayhi2(name) = println("Привет $(name)! Как твои дела?")
```

```julia
f2(x) = x^3
```

```julia
sayhi2("R2D2")
```

```julia
f2(42)
```

Последний способ заключается в использовании "анонимных" функций

```julia
sayhi3 = name -> println("Всем привет! И тебе, $(name), тоже!")
```

```julia
f3 = x -> x^4
```

```julia
sayhi3("Чуи")
```

```julia
f3(42)
```

## Утиная типизация в Julia
*"Если это крякает как утка, то это утка."* <br><br>
Функции Julia будут работать с любыми входными данными, с которыми их определение имеет смысл. <br><br>
Например, `sayhi` работает точно также, если в качестве параметра `name` передать число

```julia
sayhi(55595472)
```

А функция `f` будет работать для квадратных матриц. 

```julia
A = rand(3, 3)
A
```

```julia
f(A)
```

`f` так же будет работать, если передать строку "hi", поскольку для строковых переменных определена операция `^`, возвращающая копирование строки

```julia
f("hi")
```

С другой стороны, `f` не будет работать с векторами. В отличии от операции `A^2`, которая имеет смысл, операция возведения в степень вектора `v^2` не имеет смысла с математической точки зрения.

```julia
v = rand(3)
```

```julia
# This won't work
f(v)
```

## Мутирующие и немутирующие функции

Согласно договоренности, функции, название которых оканчивается на `!` изменяют параметры, которые в них передаются, а функции без `!` - не меняют.

Например, давайте посмотрим на различие между `sort` и `sort!`.


```julia
v = [3, 5, 2]
```

```julia
sort(v)
```

```julia
v
```

`sort(v)` возвращает отсортированный вектор `v`, однако `v` остается неизменным. <br><br>

С другой сторооны, когда мы вызываем `sort!(v)`, содержимое `v` сортируется внутри переменной.

```julia
sort!(v)
```

```julia
v
```

## Некоторые высокоуровневые функции

### map

`map` - это "высокоуровневая" функция в Julia, которая *принимает функцию* в качестве одного из аргументов. 
`map` применяет данную функцию к каждому элементу структуры, которую вы также передаете в функцию. Например, выполнив

```julia
map(f, [1, 2, 3])
```
вы получите массив, где функция `f` была применена последовательно к `[1, 2, 3]`
```julia
[f(1), f(2), f(3)]
```

```julia
map(f, [1, 2, 3])
```

Здесь мы получили квадраты элементов массива `[1, 2, 3]`, а не "квадрат" массива `[1, 2, 3]`.

В случае, если у нас нет в наличии необходимой функции что бы передать ее в `map`, мы можем использовать анонимную функцию, которая создается "налету"

```julia
map(x -> x^5, [1, 2, 3])
```

И теперь у нас есть пятые степени массива `[1, 2, 3]`

### broadcast

`broadcast` другая высокоуровневая функция по-типу `map`. `broadcast` является обобщением `map`, так что она может делать все, что может `map` и даже больше. Синтаксис выова `broadcast` такой же как и у `map`

```julia
broadcast(f, [1, 2, 3])
```

и снова, мы преминяем `f` (возведение в квадрат) ко всем элементам `[1, 2, 3]` - теперь средствами "броадкастинга" `f`!

Другой способ вызова функции `broadcast` - это разместить `.` между именем функции и ее аргументами. Например,

```julia
broadcast(f, [1, 2, 3])
```
тоже самое, что и 
```julia
f.([1, 2, 3])
```

```julia
f.([1, 2, 3])
```

Снова обратите внимание, как это отличается от вызова
```julia
f([1, 2, 3])
```
Теперь мы можем возводить в квадрат элементов вектора, а не весь  вектор целиком!

И, чтобы дойти до конца, посмотрим на разницу между

```julia
f(A)
```
и
```julia
f.(A)
```
для матрицы `A`:

```julia
A = [i + 3*j for j in 0:2, i in 1:3]
```

```julia
f(A)
```

Как мы уже раньше видели, для матрицы `A`,
```
f(A) = A^2 = A * A
``` 

С другой стороны,

```julia
B = f.(A)
```

содержит квадраты всех элементов матрицы `A`.

Этот синтаксис с точкой для броадкастинга позволяет нам писать относительно сложные элементно-зависимые выражения в виде, похожем на  математические записи. К примеру, мы можем записать:

```julia
A .+ 2 .* f.(A) ./ A
```

вместо

```julia
broadcast(x -> x + 2 * f(x) / x, A)
```

хотя обе записи по делают то же самое.

### Упраженения

#### 6.1 
Напишите функцию `add_one` которая добавляет 1 ко входному значению.

```julia

```

<!-- grade_id: cell-5119b9e9623c1cb7, points: 1, grade: true, solution: false -->
```julia
@assert add_one(1) == 2
```

<!-- grade_id: cell-50f83d27187a2064, points: 1, grade: true, solution: false -->
```julia
@assert add_one(11) == 12
```

#### 6.2 
Используйте `map` или `broadcast`, чтобы увеличить каждый элемент матрицы `A` на `1` и присвоить результат в переменную `A1`.

```julia

```

```julia
@assert A1 == [2 3 4; 5 6 7; 8 9 10]
```

#### 6.3 
Используйте синтаксис `broadcast` с точкой, чтобы увеличить каждый элемент матрицы `A1` на `1` и поместить его в переменную `A2`

```julia

```

<!-- grade_id: cell-f3bd5479679a8fe1, points: 0, grade: true, solution: false -->
```julia
@assert A2 == [3 4 5; 6 7 8; 9 10 11]
```

ЗАДАНИЕ 1

```julia
function getdigits(a)
    chisla = [];
    num = a;
    while (num/10) > 1
        push!(chisla, num%10);
        num = num÷10;
    end
    push!(chisla, num);
    return chisla[end:-1:1];
end

getdigits(658)


```

Output:
```
3-element Vector{Any}:
 6
 5
 8
```

```julia


function touppercase(str_)
    s=""
    println(s)
    for simv in str_
        if 'a'<=simv<='z'
             nsimv=Char(Int(simv)-Int('a')+Int('A'))
             s *= nsimv
        else
            s *= simv
        end
    end
    return s 
end

touppercase("straNiCa")
```

Output:
```
"STRANICA"
```

```julia
function touppercase(str_)
    buf = [];
    for sym in str_
        if (sym>='a')&&(sym<='z')
            push!(buf, sym + ('A' - 'a'));
        else
            push!(buf, sym);
        end
    end
    return join(buf);
end

touppercase("STraNu")
```

Output:
```
"STRANU"
```

```julia
function isleap(year)
    if year%400==0 return true
    elseif year%100==0 return false
    elseif year%4==0 return true
    else
        return false
    end
end

isleap(2024)
```

Output:
```
true
```

```julia
function chessfloor(cell1, cell2)
    num1=Int(cell1[1])-Int('a')
    num2=Int(cell2[1]-Int('a'))

if ((cell1[2]+num1)%2==(cell2[2]+num2)%2) return true
else return false
end
end

chessfloor(['a',1],['a',3])
```

Output:
```
true
```

```julia


function firstnorm(vec_::AbstractVector{<:Number})
    norm=0
    for i in vec_
     norm += abs(i)
    end
    return norm
 end

 vec=1:1000
 firstnorm(vec)
```

Output:
```
500500
```

```julia
function infnorm(mat_::AbstractMatrix{<:Number})
    rows, cols = size(mat_)
    for j=1:cols
        for i=1:rows
            mat_[i,j]=(mat_[i,j])^2
        end
    end
    sum=0
    for j=1:cols
        for i=1:rows
            sum+=mat_[i,j]
        end
    end
    return (sum)^(1/2)
end

b = rand(10, 10)
infnorm(b)

```

Output:
```
5.283524110466422
```